# Dummy Assignment for Testing Automatic Grading

The purpose of this assignment is to get you familiar with the automatic grading system. Make sure you only modify the section marked as
```python
###
### YOUR CODE HERE
###
```
Do not modify any code before or after the cell containing the required your code or this could raise an error. You are allowed to add more cells to the notebook if needed. Press the **Submit** button on the top when you are ready to get the assignment graded. The grading takes several seconds. We recommend you make a copy of the assignment before making any changes in case you make changes to undesired sections by mistake. You can also download the assignment and work on it on your local system before submitting online.

### Q1. Complete the function so it returns the addition of the inputs [1 pt]

```python
def sum(x, y):
    ###
    ### YOUR CODE HERE
    ###
    val = x + y
    return val


# Just an example of what it should work
x = 2
y = 3

print("Printing an example:")
print("  {} + {} = {}".format(x, y, sum(x, y)))
```

Output:
```
Printing an example:
  2 + 3 = 5
```

The following tests are worth [1 pt]

<!-- grade_id: Q1, points: 1, grade: true, solution: false -->
```python
# Checking that the function is correct. Note that we have are showing you some tests but keep some
# some hidden. Try hardcoding the function so it satisfies the visible test. It should fail the
# hidden ones.

assert sum(2, 3) == 5

###
### AUTOGRADER TEST - DO NOT REMOVE
###
```

### Q2. Assign a value to the variable `x` so it is within 0.01 from 2.71 [1 pt]

```python
# Initializing the variable
x = 2.70

###
### YOUR CODE HERE
###
```

The following tests are worth [1 pt]

<!-- grade_id: Q2, points: 1, grade: true, solution: false -->
```python
# Check what happens if you set x to be 2.71, 2.719 and 2.72

###
### AUTOGRADER TEST - DO NOT REMOVE
###
```