# Vector Calculus Assignment

Answer the following questions by completing the corresponding scripts. Make sure you don't use any more packages than the ones provided here.

```python
import numpy as np


# This is the function that we will be using to check accuracy
def AlmostEqual(P, Q, digits):
    epsilon = 10**-digits
    return np.linalg.norm(P - Q) < epsilon
```

## Problem 1 [4 pts]

For parts (a) and (b), calculate the gradient of $f$ with respect to its input, i.e., compute $Df$.


### Part (a)

$f(x,y,z) = 4x^3y^2-e^zy^4+\frac{z^3}{x^2}+4y-x^{16}$

Complete the code below.

```python
def Df(x, y, z):
    # Initializing the values
    f_x = 0
    f_y = 0
    f_z = 0

    ###
    ### YOUR CODE HERE
    ###

    f_x = (12 * (x**2) * (y**2)) - ((2 * z**3 / (x**3))) - (16 * (x**15))
    f_y = 8 * y * (x**3) - (4 * (y**3) * (np.e**z)) + 4
    f_z = -(np.e**z) * (y**4) + (3 * (z**2) / (x**2))

    return np.array([[f_x, f_y, f_z]])
```

Running first test. [1 pt]

<!-- grade_id: P1a Test 1, points: 1, grade: true, solution: false -->
```python
# Test input
x = 2
y = 4
z = -1

# True value of gradient for test input
Df_true = np.array([-5.23519750e05, 1.65822863e02, -9.34271369e01])

# Check for correctness of result
assert AlmostEqual(Df_true, Df(x, y, z), 3)
```

Running second test. [1 pt]

<!-- grade_id: P1a Test 2, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
```

### Part (b)

$f\left( {u,v,p,t} \right) = 8{u^2}{t^3}p - \sqrt v \,{p^2}{t^{ - 5}} + 2{u^2}t + 3{p^4} - v$

Complete the code below.

```python
def Df(u, v, p, t):
    # Initializing the values
    f_u = 0
    f_v = 0
    f_p = 0
    f_t = 0

    ###
    ### YOUR CODE HERE
    ###
    f_u = 16 * u * t**3 * p + 4 * u * t
    f_v = -(p**2 * t**-5) / (2 * np.sqrt(v)) - 1
    f_p = 8 * u**2 * t**3 - 2 * p * np.sqrt(v) * t**-5 + 12 * p**3
    f_t = 24 * u**2 * t**2 * p + 5 * t**-6 * np.sqrt(v) * p**2 + 2 * u**2

    return np.array([[f_u, f_v, f_p, f_t]])
```

Running first test. [1pt]

<!-- grade_id: P1b Test 1, points: 1, grade: true, solution: false -->
```python
# Test input
u = -1
v = 1
p = 2
t = 1

# True value of gradient for test input
Df_true = np.array([-36, -3, 100, 70])

# Check for correctness of result
assert AlmostEqual(Df_true, Df(u, v, p, t), 3)
```

Running second test. [1pt]

<!-- grade_id: P1b Test 2, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
```

## Problem 2 [6 pts]

Using the chain rule, compute the derivative $\frac{df\circ u(x)}{dx}$ where $u:\mathbb{R}^2\to\mathbb{R}^2$ and $f:\mathbb{R}^2\to\mathbb{R}$ given that
$f(z) = (z_1-z_2)e^{z_1}$ and $u(x) = \begin{bmatrix} x_1 x_2 \\ x_1^2 - x_2^2 \end{bmatrix}$. Complete the code for each of the following parts:

a. Compute $Df = \frac{df(z)}{dz}$

b. Compute $Du = \frac{du(x)}{dx}$

c. Use the chain rule and the values of $Df, Du, f, u$, and $x$ to compute $Df\circ u[x]$.

```python
def f(z1, z2):
    return (z1 - z2) * np.exp(z1)


def u(x1, x2):
    return np.array([[x1 * x2], [x1**2 - x2**2]])
```

Complete the code below.

```python
# Part (a)
# Remember that Df is a row vector in this case.
def Df(z_1, z_2):
    ###
    ### YOUR CODE HERE
    ###
    df_dz1 = (1 + z_1 - z_2) * np.exp(z_1)
    df_dz2 = -np.exp(z_1)
    return np.array([df_dz1, df_dz2])


# Part (b)
# Remember that Du is a matrix in this case.
def Du(x_1, x_2):
    ###
    ### YOUR CODE HERE
    ###

    du_dx = np.array([[x_2, x_1], [2 * x_1, -2 * x_2]])
    return du_dx


# Part (c)
# You should only use x_1,x_2 and the function u, Du and Df together
# with matrix multiplications.
def Dfou(x_1, x_2):
    ###
    ### YOUR CODE HERE
    ###

    u_val = u(x_1, x_2)
    Df_val = Df(u_val[0, 0], u_val[1, 0])
    Du_val = Du(x_1, x_2)
    return Df_val @ Du_val
```

Running tests for Df. [2 pt]

<!-- grade_id: P2a Test 1, points: 1, grade: true, solution: false -->
```python
# Part (a)

# Test input
z_1 = 0
z_2 = 1

# True value of Df for test input
Df_true = np.array([[0, -1]])

# Check for correctness of result
assert AlmostEqual(Df_true, Df(z_1, z_2), 3)
```

<!-- grade_id: P2a Test 2, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
```

Running tests for Du. [2 pt]

<!-- grade_id: P2b Test 1, points: 1, grade: true, solution: false -->
```python
# Part (b)

# Test input
x_1 = 1
x_2 = 1

# True value of Du for test input
Du_true = np.array([[1, 1], [2, -2]])

# Check for correctness of result
assert AlmostEqual(Du_true, Du(x_1, x_2), 3)
```

<!-- grade_id: P2b Test 2, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
```

Running tests for Dfou. [2 pts]

<!-- grade_id: P2c Test 1, points: 1, grade: true, solution: false -->
```python
# Part (c)

# Test input
x_1, x_2 = 0, 1

# True value of dfou for test input
Dfou_true = np.array([[2.0, 2.0]])

# Check for correctness of result
assert AlmostEqual(Dfou_true, Dfou(x_1, x_2), 3)
```

<!-- grade_id: P2c Test 2, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
```

## Problem 3 [2 pts]

If $A(t) = \begin{bmatrix}
t^2 && t+1 \\
t^3+t+3 &&  7
\end{bmatrix}$, calculate $DA^2[t] = \frac{dA^2(t)}{dt}$.

Complete the code below.

```python
def DA2(t):
    ###
    ### YOUR CODE HERE
    ###

    def A(t):
        return np.array([[t**2, t + 1], [t**3 + t + 3, 7]])

    def DA(t):
        return np.array([[2 * t, 1], [3 * t**2 + 1, 0]])

    At = A(t)
    dAt = DA(t)
    dA2dt = dAt @ At + At @ dAt

    return dA2dt
```

Running first test. [1 pt]

<!-- grade_id: P3 Test 1, points: 1, grade: true, solution: false -->
```python
# Test input
t = 0

# True value of dA2dt for test input
DA2_true = np.array([[4, 7], [7, 4]])

# Checking for correctness of result
assert AlmostEqual(DA2_true, DA2(t), 3)
```

Running second test. [1 pt]

<!-- grade_id: P3 Test 2, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
```

## Problem 4 [8 pts]

Using the chain rule, compute the derivative $\frac{df\circ g \circ h(x)}{dx}$ where $h:\mathbb{R}^3\to\mathbb{R}^2$, $g:\mathbb{R}^2\to\mathbb{R}^2$, and $f:\mathbb{R}^2\to\mathbb{R}$ given $f(z)=z_1 z_2+1$, $g(y)= \begin{bmatrix}
y_1 + 2y_2 \\
4y_1 - y_2
\end{bmatrix}$, and  $h(x) =  \begin{bmatrix}
e^{x_1}\cos x_2 + x_3 \\
e^{x_1}\sin x_2 + x_3
\end{bmatrix} $. Complete the code for each of the following parts:

a. Compute $Df = \frac{df(z)}{dz}$

b. Compute $Dg = \frac{dg(y)}{dy}$

c. Compute $Dh = \frac{dh(x)}{dx}$

d. Use the chain rule and the values of $Df, Dg, Dh, f, g, h$ and $x$ to compute $Df\circ g \circ h[x] $

```python
def f(z_1, z_2):
    return z_1 * z_2 + 1


def g(y_1, y_2):
    return np.array([[y_1 + 2 * y_2], [4 * y_1 - y_2]])


def h(x_1, x_2, x_3):
    return np.array(
        [[np.exp(x_1) * np.cos(x_2) + x_3], [np.exp(x_1) * np.sin(x_2) + x_3]]
    )
```

Complete the code below.

```python
# Part (a)
def Df(z_1, z_2):
    ###
    ### YOUR CODE HERE
    ###
    dz1 = z_2
    dz2 = z_1

    return np.array([dz1, dz2])


# Part (b)
def Dg(y_1, y_2):
    ###
    ### YOUR CODE HERE
    ###

    dy1 = [1, 2]
    dy2 = [4, -1]
    return np.array([dy1, dy2])


# Part (c)
def Dh(x_1, x_2, x_3):
    ###
    ### YOUR CODE HERE
    ###
    dz1 = [np.e**x_1 * np.cos(x_2), np.e**x_1 * np.sin(x_2)]
    dz2 = [-np.e**x_1 * np.sin(x_2), np.e**x_1 * np.cos(x_2)]
    dz3 = [1, 1]
    return np.stack([dz1, dz2, dz3], axis=1).squeeze()


# Part (d)
# Make sure to only use x_1, x_2, x_3 and the functions f, g, h, Df, Dg, Dh
# together with simple matrix multiplications.
def Dfogoh(x_1, x_2, x_3):
    ###
    ### YOUR CODE HERE
    ###
    hx = h(x_1, x_2, x_3)
    ghx = g(hx[0], hx[1]).squeeze(-1)

    df = Df(ghx[0], ghx[1]).squeeze(-1)
    dg = Dg(hx[0], hx[1])
    dh = Dh(x_1, x_2, x_3)

    return df.T @ dg @ dh
```

Running tests for Df. [2 pts]

<!-- grade_id: P4a Test 1, points: 1, grade: true, solution: false -->
```python
# Part (a)

z_1, z_2 = 1, 2

Df_true = np.array([[2, 1]])

assert AlmostEqual(Df_true, Df(z_1, z_2), 3)
```

<!-- grade_id: P4a Test 2, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
```

Running tests for Dg. [2 pts]

<!-- grade_id: P4b Test 1, points: 1, grade: true, solution: false -->
```python
# Part (b)

y_1, y_2 = 1, 2

Dg_true = np.array([[1, 2], [4, -1]])

assert AlmostEqual(Dg_true, Dg(y_1, y_2), 3)
```

<!-- grade_id: P4b Test 2, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
```

Running tests for Dh. [2 pts]

<!-- grade_id: P4c Test 1, points: 1, grade: true, solution: false -->
```python
# Part (c)

x_1, x_2, x_3 = 0, 0, 0

Dh_true = np.array([[1, 0, 1], [0, 1, 1]])

assert AlmostEqual(Dh_true, Dh(x_1, x_2, x_3), 3)
```

<!-- grade_id: P4c Test 2, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
```

Running tests for Dfogoh. [2 pts]

<!-- grade_id: P4d Test 1, points: 1, grade: true, solution: false -->
```python
# Part (d)

x_1, x_2, x_3 = 0, 0, 0

Dfogoh_true = np.array([[8, 7, 15]])

assert AlmostEqual(Dfogoh_true, Dfogoh(x_1, x_2, x_3), 3)
```

<!-- grade_id: P4d Test 2, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###
```