# Backpropagation Assignment

Be sure to follow the guidelines below.
* Do not use any more packages than the ones provided in the notebook.
* Do not make any changes outside the blocks that state "YOUR CODE HERE".
* To make sure that your assignment is submitted correctly, click the "Submit" button and check the grading report.

In this assignment, you will implement the steps for backpropagation on the simple computation graph shown below.

![Computation Graph](computation_graph.png "Computation Graph")

Note that we start the variables with the index 0 in order to better match how Python indexes things. The relationships between different variables are defined as follows:
$$u_2 = u_0.u_1$$
$$u_3 = cos(u_2)$$
$$u_4 = sin(u_2)$$
$$u_5 = σ(u_3)$$
$$u_6 = u_4^2 + u_5^2$$
where $u_0$ and $u_1$ are the inputs, and $\sigma$ is the sigmoid function. These functions and their derivatives based on their input are all implemented below.

```python
import numpy as np


def sigmoid_f(x):
    z = 1 / (1 + np.exp(-x))
    return z


# MAIN FUNCTIONS


def u2(uList):
    return uList[0] * uList[1]


def u3(uList):
    return np.cos(uList[2])


def u4(uList):
    return np.sin(uList[2])


def u5(uList):
    return sigmoid_f(uList[3])


def u6(uList):
    return uList[4] ** 2 + uList[5] ** 2


# DERIVATIVES


def du6u5(uList):
    val = 2 * uList[5]
    return val


def du6u4(uList):
    val = 2 * uList[4]
    return val


def du5u3(uList):
    val = sigmoid_f(uList[3]) * (1 - sigmoid_f(uList[3]))
    return val


def du3u2(uList):
    val = -np.sin(uList[2])
    return val


def du4u2(uList):
    val = np.cos(uList[2])
    return val


def du2u1(uList):
    val = uList[0]
    return val


def du2u0(uList):
    val = uList[1]
    return val
```

Below you can see an illustration of how these functions can be used.

```python
# Performing a pass of forward propagation
uList = np.zeros(7)
uList[0:2] = [1, 5]
uList[2] = u2(uList)
uList[3] = u3(uList)
uList[4] = u4(uList)
uList[5] = u5(uList)
uList[6] = u6(uList)

for k in range(len(uList)):
    print("u{:d}: {:6.3f}".format(k, uList[k]))
```

Output:
```
u0:  1.000
u1:  5.000
u2:  5.000
u3:  0.284
u4: -0.959
u5:  0.570
u6:  1.245
```

The following lines of code illustrate how to call the derivatives: $\frac{du_2}{du_0}(u_0,u_1), \frac{du_5}{du_3}(u_3)$. The rest of the derivatives can be obtained in a similar manner.

```python
print("du2/du0: {:6.3f}".format(du2u0(uList)))
print("du5/du3: {:6.3f}".format(du5u3(uList)))
```

Output:
```
du2/du0:  5.000
du5/du3:  0.245
```

### Perform Backpropagation

In this section, you need to follow the steps of the backpropagation algorithm to implement the derivatives $B_i = \frac{du_6}{du_i}$ below. Only modify the blocks of code marked. Note that you should be using the values of $B_i$ pre-computed in prior stages and the corresponding derivatives so your implementation of each function should be one line long.

```python
def B5(BList, uList):
    ###
    ### YOUR CODE HERE
    ###
    return BList[6] * du6u5(uList)


def B4(BList, uList):
    ###
    ### YOUR CODE HERE
    ###
    return BList[6] * du6u4(uList)


def B3(BList, uList):
    ###
    ### YOUR CODE HERE
    ###

    return BList[5] * du5u3(uList)


def B2(BList, uList):
    ###
    ### YOUR CODE HERE
    ###

    return BList[3] * du3u2(uList) + BList[4] * du4u2(uList)


def B1(BList, uList):
    ###
    ### YOUR CODE HERE
    ###

    return BList[2] * du2u1(uList)


def B0(BList, uList):
    ###
    ### YOUR CODE HERE
    ###

    return BList[2] * du2u0(uList)
```

Below you can see an illustration of how these functions can be used.

```python
# Performing a pass of backpropagation
BList = np.zeros(7)
BList[6] = 1
BList[5] = B5(BList, uList)
BList[4] = B4(BList, uList)
BList[3] = B3(BList, uList)
BList[2] = B2(BList, uList)
BList[1] = B1(BList, uList)
BList[0] = B0(BList, uList)

for k in range(len(BList)):
    print("B{:d}: {:6.3f}".format(k, uList[k]))
```

Output:
```
B0:  1.000
B1:  5.000
B2:  5.000
B3:  0.284
B4: -0.959
B5:  0.570
B6:  1.245
```

### Tests

Now we run tests to check if your solutions are correct.

```python
# This is the function that we will be using to check accuracy
def AlmostEqual(P, Q, digits):
    epsilon = 10**-digits
    return np.linalg.norm(P - Q) < epsilon
```

Running tests for B5 [2pt]

<!-- grade_id: B5, points: 2, grade: true, solution: false -->
```python
# FIRST TEST

uList = np.zeros(7)
uList[0:2] = [2, 5]
uList[2] = u2(uList)
uList[3] = u3(uList)
uList[4] = u4(uList)
uList[5] = u5(uList)
uList[6] = u6(uList)

BList = np.array([3, 5, 6, 2, 1, 7, 2])

B5_true = np.array([1.2069214669243271])

assert AlmostEqual(B5_true, B5(BList, uList), 3)


# SECOND TEST

###
### AUTOGRADER TEST - DO NOT REMOVE
###
```

Running tests for B4 [2pt]

<!-- grade_id: B4, points: 2, grade: true, solution: false -->
```python
# FIRST TEST

uList = np.zeros(7)
uList[0:2] = [2, 5]
uList[2] = u2(uList)
uList[3] = u3(uList)
uList[4] = u4(uList)
uList[5] = u5(uList)
uList[6] = u6(uList)

BList = np.array([3, 5, 6, 2, 1, 7, 2])

B4_true = np.array([-2.176084443557479])

assert AlmostEqual(B4_true, B4(BList, uList), 3)


# SECOND TEST

###
### AUTOGRADER TEST - DO NOT REMOVE
###
```

Running tests for B3 [2pt]

<!-- grade_id: B3, points: 2, grade: true, solution: false -->
```python
# FIRST TEST

uList = np.zeros(7)
uList[0:2] = [2, 5]
uList[2] = u2(uList)
uList[3] = u3(uList)
uList[4] = u4(uList)
uList[5] = u5(uList)
uList[6] = u6(uList)

BList = np.array([3, 5, 6, 2, 1, 7, 2])

B3_true = np.array([1.4748240676638606])

assert AlmostEqual(B3_true, B3(BList, uList), 3)


# SECOND TEST

###
### AUTOGRADER TEST - DO NOT REMOVE
###
```

Running tests for B2 [2pt]

<!-- grade_id: B2, points: 2, grade: true, solution: false -->
```python
# FIRST TEST

uList = np.zeros(7)
uList[0:2] = [2, 5]
uList[2] = u2(uList)
uList[3] = u3(uList)
uList[4] = u4(uList)
uList[5] = u5(uList)
uList[6] = u6(uList)

BList = np.array([3, 5, 6, 2, 1, 7, 2])

B2_true = np.array([0.2489706927022871])

assert AlmostEqual(B2_true, B2(BList, uList), 3)


# SECOND TEST

###
### AUTOGRADER TEST - DO NOT REMOVE
###
```

Running tests for B1 [2pt]

<!-- grade_id: B1, points: 2, grade: true, solution: false -->
```python
# FIRST TEST

uList = np.zeros(7)
uList[0:2] = [2, 5]
uList[2] = u2(uList)
uList[3] = u3(uList)
uList[4] = u4(uList)
uList[5] = u5(uList)
uList[6] = u6(uList)

BList = np.array([3, 5, 6, 2, 1, 7, 2])

B1_true = np.array([12.0])

assert AlmostEqual(B1_true, B1(BList, uList), 3)


# SECOND TEST

###
### AUTOGRADER TEST - DO NOT REMOVE
###
```

Running tests for B0 [2pt]

<!-- grade_id: B0, points: 2, grade: true, solution: false -->
```python
# FIRST TEST

uList = np.zeros(7)
uList[0:2] = [2, 5]
uList[2] = u2(uList)
uList[3] = u3(uList)
uList[4] = u4(uList)
uList[5] = u5(uList)
uList[6] = u6(uList)

BList = np.array([3, 5, 6, 2, 1, 7, 2])

B0_true = np.array([30.0])

assert AlmostEqual(B0_true, B0(BList, uList), 3)


# SECOND TEST

###
### AUTOGRADER TEST - DO NOT REMOVE
###
```